*Adapté d'un [sujet d'épreuve pratique des ENS](http://informatique.ens-lyon.fr/concours-info/2012/sujet-2012-6.pdf) de 2012*

On considère un ensemble **fini** $E$, une application $f : E \to E$.

À partir d'un élément $x_0 \in E$, on peut construire une suite $\left(x_{n}\right)_{n \in \mathbb{N}}$, appelée *suite des itérés* de $x_0$, telle que pour tout $n \in \mathbb{N}$, $x_{n+1} = f(x_n)$.

Comme l'ensemble $E$ est fini, la suite $\left(x_{n}\right)_{n \in \mathbb{N}}$ prend un nombre fini de valeurs, donc il existe $i < j$ tel que $x_i = x_j$.

Les valeurs prises par la suite à partir du rang $i$ seront alors : $x_i, x_{i+1}, \dots, x_{j-1}, x_i, x_{i+1}, \dots $. Autrement dit, à partir du rang $i$, la suite sera périodique de période $\lambda = j-i$.

On appelera *pré-période* l'indice minimal $\mu = i$ tel qu'il existe $j > i$ tel que $x_j = x_i$ (autrement dit le nombre d'éléments de la suite n'appartenant pas au cycle).

<font size="5">👩‍💻</font> En utilisant le fait que pour tout $n \in \mathbb{N}$, $x_{n+1}  =f(x_n)$, écrire une fonction `itere : ('a -> 'a) -> 'a -> int -> 'a` qui prend en argument la fonction $f$, l'élément $x_0$ et l'entier $n$, et qui renvoie la valeur de $x_n$.

<!-- grade_id: cell-e118861d4f77105e, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-c2527ef79ede70f1, points: 0, grade: true, solution: false -->
```OCaml
(*Exécuter cette cellule pour tester votre fonction *)
let () = assert (itere (fun x -> (2*x) mod 2048) 1 10 = 1024);;
let () = assert (itere (fun x -> (x*x + 17) mod 431) 4 2000 = 46);;
let () = assert (itere (fun x -> (x*x + 92) mod 32069) 33 321 = 2368);;
```

<font size="5">👨🏽‍💻 </font>  En remarquant que pour tout $n \in \mathbb{N}^*$, $x_n$ peut aussi s'obtenir en appliquant $n-1$ fois la fonction $f$ à $f(x_0)$, en déduire une fonction `itere2` qui calcule différemment la valeur de $x_n$.

<!-- grade_id: cell-ed1704fc45567b36, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-ec9bfed394b8d014, points: 0, grade: true, solution: false -->
```OCaml
(*Exécuter cette cellule pour tester votre fonction *)
let () = assert (itere2 (fun x -> (2*x) mod 2048) 1 10 = 1024);;
let () = assert (itere2 (fun x -> (x*x + 17) mod 431) 4 2000 = 46);;
let () = assert (itere2 (fun x -> (x*x + 234) mod 32069) 33 321 = 18149);;
```

<font size="5">🖊️</font> On pourra remarquer la différence de comportement de ces deux fonctions sur les appels suivants :

```OCaml
itere (fun x -> (x*x + 17) mod 431) 4 2000000;;
```

```OCaml
itere2 (fun x -> (x*x + 17) mod 431) 4 2000000;;
```

# [Algorithme du lièvre et de la tortue](https://fr.wikipedia.org/wiki/Algorithme_du_li%C3%A8vre_et_de_la_tortue)

On utilisera dans cette partie l'agorithme de détection de cycle de Floyd, aussi appelé algorithme du lièvre et de la tortue.

Pour cela, on considère la suite  $\left(y_{n}\right)_{n \in \mathbb{N}}$ définie par $y_0 = x_0$ et pour tout $n \in \mathbb{N}$, $y_{n+1} = f(f(y_n))$.

<font size="5">👩🏿‍💻</font>  Écrire une fonction `floyd` qui prend en argument la fonction $f$ et l'élément $x_0$, et qui renvoie un couple constitué du plus petit entier strictement positif $i$ tel que $x_i = y_i$ et de la valeur $x_i$ correspondante.

On pourra utiliser une fonction récursive auxilliaire.


<!-- grade_id: cell-20016cda4dada42b, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

```OCaml
(*Exécuter cette cellule pour tester votre fonction *)
let () = assert (floyd (fun x -> (2*x) mod 431) 1  = (43, 1));;
let () = assert (floyd (fun x -> (x*x + 234) mod 32069) 33 = (95, 10002));;
let () = assert (floyd (fun x -> (x*x + 92) mod 32069) 33 = (320, 11326));;
let () = assert (floyd (fun x -> (x*x + 17) mod 431) 4 = (25, 46));;
```

<font size="5">👨🏼‍💻</font> Montrer que la valeur de $i$ calculée correspond au plus petit multiple de la période $\lambda$ supérieur ou égal à la pré-période $\mu$.

<!-- grade_id: cell-21c9632765159855, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">👩🏻‍💻</font>  En remplaçant $x_0$ par la valeur de $x_i$ ainsi obtenue, quelle sera la nouvelle valeur de $i$ calculée par la fonction `floyd` ?

<!-- grade_id: cell-967b7c1aca5a643c, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">👨🏿‍💻</font>  En déduire une fonction `periode` qui prend en argument $f$ et $x_0$ et qui renvoie la période $\lambda$ de la suite des itérés $\left(x_{n}\right)_{n \in \mathbb{N}}$

<!-- grade_id: cell-ffa65766d304a9b5, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-ebb25c8cfa8c46ff, points: 0, grade: true, solution: false -->
```OCaml
(*Exécuter cette cellule pour tester votre fonction *)
let () = assert (periode (fun x -> (2*x) mod 431) 1  = 43);;
let () = assert (periode (fun x -> (x*x + 234) mod 32069) 33 = 95);;
let () = assert (periode (fun x -> (x*x + 92) mod 32069) 33 = 8);;
let () = assert (periode (fun x -> (x*x + 17) mod 431) 4 = 25);;
```

<font size="5">👩‍💻</font>  Montrer que $x_{i+\mu} = x_{\mu}$ 

<!-- grade_id: cell-b9bd9d6421de397d, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">👨🏾‍💻</font>  En déduire une fonction `pre_periode` qui prend en argument $f$ et $x_0$ et qui renvoie la pré-période $\mu$ de la suite $\left(x_{n}\right)_{n \in \mathbb{N}}$.

<!-- grade_id: cell-207412f53fd66597, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-0b824541da6ecd45, points: 0, grade: true, solution: false -->
```OCaml
(*Exécuter cette cellule pour tester votre fonction *)
let () = assert (pre_periode (fun x -> (2*x) mod 431) 1  = 0);;
let () = assert (pre_periode (fun x -> (x*x + 234) mod 32069) 33 = 28);;
let () = assert (pre_periode (fun x -> (x*x + 17) mod 431) 4 = 5);;
```

# Algorithme de Brent
L'algorithme de Floyd permet de trouver une valeur $x_{i}$ dans le cycle et ensuite la longueur du cycle et la longueur de la pré-période en considérant la suite d'indices $(i, 2 i)$ et en testant l'égalité $x_{i}=x_{2 i} .$

L'algorithme de Brent utilise la suite $(i, j)$ et teste l'égalité $x_{i}=x_{j} .$ La suite $(i, j)$ est construire en partant de $(0,1)$ et telle que
$$
\begin{aligned}
(i, j) & \mapsto(i, j+1) \quad \text { si } \quad j \leq 2 i\\
(i, j) & \mapsto(j, j+1) \quad \text { si } \quad j = 2 i
\end{aligned} 
$$
<font size="5">👨🏻‍💻 </font> Écrire une fonction `brent` qui prend en argument $f$ et $x_0$ et qui renvoie le couple $(i,j)$ obtenu par cet algorithme.

<!-- grade_id: cell-73747b8096e1499d, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-750da941a51cc2d8, points: 0, grade: true, solution: false -->
```OCaml
(*Exécuter cette cellule pour tester votre fonction *)
let () = assert (brent (fun x -> (2*x) mod 431) 1  = (63, 106));;
let () = assert (brent (fun x -> (x*x + 234) mod 32069) 33  = (127, 222));;
```