# Pour démarrer

Vous vous trouvez actuellement dans un *notebook* Jupyter. Celui-ci contient deux types de cellule : des cellules de texte (comme celle-ci), et des cellules de code.

Les cellules peuvent être exécutées via le menu **Run**, via l'icône Play ou avec un raccourci clavier :
- `Shift+Entrée` pour exécuter la cellule et passer à la suivante ;
- `Ctrl+Entrée` pour exécuter la cellule sans passer à la suivante ;
- `Alt+Entrée` pour exécuter la cellule et créer une nouvelle cellule en dessous.

Par exemple, la cellule de code suivante définit une fonction permettant de calculer la factorielle d'un entier naturel.

```OCaml
let rec fact n =
    match n with
    | 0 -> 1
    | _ -> n * fact (n-1)
;;
```

Output:
```
val fact : int -> int = <fun>
```

La cellule ci-dessous appelle la fonction `fact` et ne pourra donc être exécutée sans erreur que si la cellule du dessus a été exécutée.

```OCaml
fact 6;;
```

Output:
```
- : int = 720
```

Par défaut, les nouvelles cellules créées sont des cellules de code. Pour modifier le style d'une cellule, passez en mode commande à l'aide de la touche `Echap` si vous êtes en mode édition, puis tapez `m` pour transformer la cellule courante en cellule de texte ou `y` pour la transformer en cellule de code.

Vous pouvez par ailleurs créer une nouvelle cellule au-dessus de la cellule courante en tapant `A` (*above*) ou en dessous en tapant `B`.

Pour passer du mode commande au mode édition, taper sur `Entrée`.

Pour la mise en forme des cellules de texte, vous pouvez consulter l'aide sur Markdown présente dans le menu **Help**

Les cellules déjà présentes dans ce *notebook* sont de plusieurs types :
- Des cellules contenant l'énoncé, qui sont en lecture seule ;
- Des cellules de réponses, que vous pouvez (et devez) modifier ;
- Des cellules permettant de tester vos réponses, qui sont en lecture seule.

Vous pouvez ajouter d'autres cellules si vous le souhaitez.

# Pour rendre le TP

Les TP ne seront pas à rendre systématiquement. Lorsqu'un TP devra être rendu, cela sera indiqué dans le *notebook*.

**Ce premier TP est à rendre en traitant (au moins) les questions de programmation.** Les questions portant sur la complexité des fonctions sont facultatives.

Pour cela, il faudra déposer votre *notebook* au nom `TP01.ipynb` dans les deux semaines suivant le TP dans le dossier `Depot_fichiers` présent [à l'adresse suivante](https://cloud.maths.mlong.fr/), en vous connectant avec votre adresse mail et le mot de passe que vous aurez défini (utiliser *mot de passe oublié* pour la première connexion).

# Quelques fonctions autour des dates

#### Règles de validité d'une date :
- janvier, mars, mai, juillet, août, octobre, décembre comportent 31 jours ; avril, juin, septembre, novembre comportent 30 jours ; février comporte  29 jours si l'année est bissextile, 28 sinon ;
- une année est bissextile
  - pour les années séculaires, si elle est divisible par 400
  - pour les autres années, si elle est divisible par 4

<font size="5">👩‍💻</font>  Écrire la fonction Caml `bissextile : int -> bool` qui renvoie `true` si l'année `n` est bissextile.

<div class="alert alert-block alert-info">
    On pourra utiliser une expression <code>if e1 then e2 else e3</code>, où <code>e1</code> est une expression booléenne. Si la valeur de <code>e1</code> est <code>true</code>, l'expression <code>e2</code> est évaluée et sa valeur est retournée ; sinon, ce sera la valeur de <code>e3</code> qui sera retournée.</div>

<!-- grade_id: cell-342f629193548472, points: 0, grade: false, solution: true -->
```OCaml
let bissextile annee =
  if annee mod 100 = 0
  then annee mod 400 = 0
  else annee mod 4 = 0
;;
```

Output:
```
val bissextile : int -> bool = <fun>
```

<!-- grade_id: cell-e41c9ca7fd9ba3d5, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (not (bissextile 2022));;
let () = assert (bissextile 2024);;
let () = assert (bissextile 2000);;
let () = assert (bissextile 2040);;
let () = assert (not (bissextile 2100));;
```

<font size="5">🧑🏿‍💻</font> La fonction suivante utilise un *filtrage*. La tester. Que renvoie-t-elle ?

```OCaml
let nb_jours m a =
  match m with
  | 2 -> if bissextile a then 29 else 28
  | 4 -> 30
  | 6 -> 30
  | 9 -> 30
  | 11 -> 30
  | _ -> 31
;;
```

Output:
```
val nb_jours : int -> int -> int = <fun>
```

<!-- grade_id: cell-f0208aaf427503bc, points: 0, grade: true, solution: true -->

Cette fonction prend en argument un entier et renvoie le nombre de jours dans le mois correspondant.


<font size="5">👨🏻‍💻</font>  Écrire une fonction `date_valide : int -> int -> int -> bool` prenant en argument les entiers `j, m , a` et renvoyant le booléen `true` si la date représentée par le triplet `j, m , a` (jour, mois, année) est une date valide, et `false` sinon.

<!-- grade_id: cell-fc2d30efe1442906, points: 0, grade: false, solution: true -->
```OCaml

let date_valide j m a =
  let nb_j = nb_jours m a in
  m > 0 && m <= 12 && j > 0 && j <= nb_j
;;

```

Output:
```
val date_valide : int -> int -> int -> bool = <fun>
```

<!-- grade_id: cell-f1dc1281d66471ef, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (date_valide 31 07 2022);;
let () = assert (not (date_valide 29 02 2022));;
let () = assert (not (date_valide 30 14 2022));;
let () = assert (not (date_valide (-2) 14 2022));;
let () = assert (not (date_valide 18 0 2022));;
(* BEGIN HIDDEN TESTS *)
let () = assert (date_valide 1 07 2022);;
let () = assert (date_valide 29 02 2024);;
let () = assert (not (date_valide 0 2 2022));;
(* END HIDDEN TESTS *)
```

			
On convient de numéroter les jours de la semaine de $0$ à $6$ (0 pour dimanche, 1 pour lundi, etc.).

<font size="5">🧑🏾‍💻</font> Écrire une fonction donnant, pour tout numéro entre $0$ et $6$, le nom du jour correspondant sous forme d'une chaîne de caractères.

<div class="alert alert-block alert-info">
    On pourra utiliser un filtrage.
</div>

<!-- grade_id: cell-b73a552ab7f16ffc, points: 0, grade: false, solution: true -->
```OCaml

let nom_jour n =
  match n with
  | 0 -> "dimanche"
  | 1 -> "lundi"
  | 2 -> "mardi"
  | 3 -> "mercredi"
  | 4 -> "jeudi"
  | 5 -> "vendredi"
  | 6 -> "samedi"
  | _ -> failwith "Argument incorrect"
;;

```

Output:
```
val nom_jour : int -> string = <fun>
```

<!-- grade_id: cell-f55cb2954afbdf55, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (nom_jour 1 = "lundi");;
let () = assert (nom_jour 2 = "mardi");;
let () = assert (nom_jour 3 = "mercredi");;
let () = assert (nom_jour 4 = "jeudi");;
let () = assert (nom_jour 5 = "vendredi");;
let () = assert (nom_jour 6 = "samedi");;
let () = assert (nom_jour 0 = "dimanche");;
```

En notant $\lfloor x \rfloor$ la partie entière de $x$, le numéro du jour de la semaine correspondant à une date donnée $(j,m,a)$ s'obtient en prenant le reste de la division par $7$ de l'entier $n$ défini par :
		$$n = j + \lfloor 2,6 p -0,2\rfloor + d + \left\lfloor \frac{d}{4} \right\rfloor +\left\lfloor \frac{c}{4} \right\rfloor +5c \qquad\text{(Formule de Zeller)}$$

où 

$$\displaystyle p = \left\lbrace \begin{array}{ll} m+10 & \text{ si } m \leqslant 2 \\
			m-2 & \text{ sinon}
		\end{array}\right. 
		\qquad b = \left\lbrace \begin{array}{ll} a-1 & \text{ si } m \leqslant 2 \\
			a & \text{ sinon}
		\end{array}\right. 
		\qquad c = \left\lfloor \frac{b}{100} \right\rfloor 
		\qquad d=b-100c $$

<font size="5">👩🏼‍💻</font>  Écrire une fonction qui calcule ce numéro à partir de $j$, $m$ et $a$.

<div class="alert alert-block alert-info">
On rappelle que OCaml est un langage fortement typé et qu'il n'est pas possible de réaliser directement une opération entre un flottant et un entier.
    
On pourra donc utiliser les fonctions <code>float_of_int</code> et <code>int_of_float</code> (qui donne la partie entière du flottant passé en paramètre).
</div>

<div class="alert alert-block alert-info">
    
Pour définir un objet dont la valeur dépend d'une condition, on pourra écrire : 
    <code>let var = if condition then e1 else e2 in</code>
</div>

<!-- grade_id: cell-78a91564a5742a3d, points: 0, grade: false, solution: true -->
```OCaml

let zeller j m a =
  let p = if  m <= 2 then m+10 else m-2 in
  let b = if  m <= 2 then a-1 else a in
  let c = b / 100 in
  let d = b - 100*c in
  let n = j
      + int_of_float (2.6 *. float_of_int p -. 0.2)
      + d
      + int_of_float (float_of_int d /. 4.)
      + int_of_float (float_of_int c /. 4.)
      + 5*c
  in n mod 7
;;

```

Output:
```
val zeller : int -> int -> int -> int = <fun>
```

<!-- grade_id: cell-115d60a767e15dfa, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (zeller 1 1 2022 = 6);;
let () = assert (zeller 17 8 1936 = 1);;
let () = assert (zeller 23 6 1912 = 0);;
```

<font size="5">🧑🏻‍💻</font> En déduire une fonction `jour_semaine` donnant le jour de la semaine correspondant à une date donnée.

<!-- grade_id: cell-c0229bc3ee524390, points: 0, grade: false, solution: true -->
```OCaml

let jour_semaine j m a =
  if date_valide j m a
  then nom_jour (zeller j m a)
  else failwith "date invalide"
;;

```

Output:
```
val jour_semaine : int -> int -> int -> string = <fun>
```

<!-- grade_id: cell-fe5383bf1a22cf26, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (jour_semaine 10 12 1815 = "dimanche");;
let () = assert (jour_semaine 26 12 1791 = "lundi");;
(* BEGIN HIDDEN TESTS *)
let () = assert (jour_semaine 1 4 1776 = "lundi");;
(* END HIDDEN TESTS *)
```

# Calculs de puissances

<font size="5">👩🏽‍💻</font> Écrire une fonction `puiss : int -> int -> int` telle que si $n$ est un entier naturel,
`puiss x n` calcule $x^{n}$ (sans utiliser `**`).

<!-- grade_id: cell-1764cba30500d0a1, points: 0, grade: false, solution: true -->
```OCaml

let rec puiss x n = 
  match n with
  | 0 -> 1
  | _ -> x * puiss x (n-1)
;;

```

Output:
```
val puiss : int -> int -> int = <fun>
```

<!-- grade_id: cell-85acdb38f2ce94da, points: 0.5, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (puiss 2 10 = 1024);;
let () = assert (puiss 2 11 = 2048);;
(* BEGIN HIDDEN TESTS *)
let () = assert (puiss 2 1 = 2);;
let () = assert (puiss 2 5 = 32);;
let () = assert (puiss 2 8 = 256);;
let () = assert (puiss 3 11 = 177147);;
(* END HIDDEN TESTS *)
```

<!-- grade_id: cell-88e1d80c595681d1, points: 0.5, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (puiss 2 0 = 1);;
(* BEGIN HIDDEN TESTS *)
let puiss x n = failwith "Doit être réécrite";;
(* END HIDDEN TESTS *)
```

<font size="5">🧑‍💻</font> Donner la complexité du calcul de `puiss x n`, en nombre de multiplications effectuées, en fonction de $n$.

<!-- grade_id: cell-50bd8a271a86a887, points: 0, grade: true, solution: true -->

<span style='color:Teal '>
Le calcul de $x^n$ demande ici $n$ multiplications.
 </span>


<font size="5">👨🏽‍💻</font> Si votre implantation de `puiss` demande plus de $10$ multiplications pour calculer $2^{42}$, récrire votre fonction.

Indication: lorsque $n$ est pair, $x^{n} = \left(x^{n/2}\right)^2$ (ou, si l'on préfère $\left(x^2\right)^{n/2}$).

<!-- grade_id: cell-a3f6a061e197eb1e, points: 0, grade: false, solution: true -->
```OCaml

let rec puiss x n = 
  if n = 0 then 1
  else
    match n mod 2 with
    | 0 -> puiss (x*x) (n/2)
    | _ -> x * puiss (x*x) (n/2)
  ;;

```

Output:
```
val puiss : int -> int -> int = <fun>
```

<!-- grade_id: cell-a21724f687584e22, points: 0.5, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (puiss 2 10 = 1024);;
let () = assert (puiss 2 11 = 2048);;
(* BEGIN HIDDEN TESTS *)
let () = assert (puiss 2 5 = 32);;
let () = assert (puiss 2 8 = 256);;
let () = assert (puiss 3 11 = 177147);;
(* END HIDDEN TESTS *)
```

<!-- grade_id: cell-0d7954c50bdcb540, points: 0.5, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (puiss 1 1000 = 1);;
```

<font size="5">🧑🏼‍💻</font> Soit $n$ un entier naturel non nul. Quelle est le nombre de multiplications effectuées dans le pire des cas lors du calcul de `puiss x n` lorsque l'entier naturel $n$ s'écrit sur $p$ bits ?

<div class="alert alert-block alert-info">
On pourra noter $u_p$ le nombre d'additions réalisées dans le pire des cas lors de l'appel <code>puiss n</code> et déterminer une relation de récurrence vérifiée par $u_p$.
</div>

En déduire la complexité dans le pire des cas de `puiss n` en fonction de $n$.

<div class="alert alert-block alert-info">    
$p$ est l'unique entier tel que $2^{p-1} \leqslant n < 2^{p}$, donc $p = \lfloor \log_2 n \rfloor + 1$. 
</div>

<!-- grade_id: cell-0a2aeed577fb7df2, points: 0, grade: true, solution: true -->


<span style='color:Teal '>
Si $p = 1$ alors $n=1$, donc $u_p = 2$.

Si $p \geqslant 2$, alors $\lfloor \frac{n}{2} \rfloor$ s'écrit sur $p-1$ bits, donc $u_{p} = 2 + u_{p-1}$. Par conséquent, $u_p = 2p$ (suite arithmétique).

Par conséquent, le calcul de $x^n$ demande un nombre de multiplication de l'ordre de $\log n$.
    </span>
    


# Suite de Fibonacci

On considère la suite de Fibonacci $(u_n)_\mathbb{N}$ définie par  $u_0=0$, $u_1=1$ et $\forall n \in \mathbb{N}, u_{n+2}=u_{n+1}+u_n$

<font size="5">👩🏻‍💻</font>  A partir de la définition, écrire une fonction `fibo` telle que `fibo n` calcule le terme $u_n$.

<!-- grade_id: cell-869b38267042bec8, points: 0, grade: false, solution: true -->
```OCaml

let rec fibo n =
  match n with
  | 0 -> 0
  | 1 -> 1
  | _ -> fibo (n-1) + fibo (n-2)
;;

```

Output:
```
val fibo : int -> int = <fun>
```

<!-- grade_id: cell-a0f6be5b4d426b50, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (fibo 2 = 1);;
let () = assert (fibo 10 = 55);;
(* BEGIN HIDDEN TESTS *)
let () = assert (fibo 0 = 0);;
let () = assert (fibo 1 = 1);;
let () = assert (fibo 15 = 610);;
(* END HIDDEN TESTS *)
```

<font size="5">👨🏿‍💻</font> Calculer quelques uns des premiers termes, puis calculer $u_{40}$. Comment expliquer que le temps de calcul soit si long ?

<!-- grade_id: cell-15ae2b5abaa63894, points: 0, grade: true, solution: true -->
```OCaml

```

<!-- grade_id: cell-7836f957bc71d8c5, points: 0, grade: true, solution: true -->


<span style='color:Teal '>
Pour calculer $u_{40}$, il faut calculer $u_{39}$ et $u_{38}$.
<br/>
Mais le calcul de $u_{39}$ demande le calcul de $u_ {38}$, qui est donc calculé deux fois, et celui de $u_{37}$, qui est aussi nécessaire à celui de $u_ {38}$.
<br/>
$u_{37}$ sera donc calculé trois fois, $u_{36}$ cinq fois (trois fois dans un appel à `fibo 37` et deux fois dans un appel à `fibo 38`), etc.
<br/>
Plus précisément, le nombre d'appels $C_n$ à la fonction `fibo` nécessaires pour calculer $u_n$ vérifie $C_0=C_1 = 1$ et $C_n = 1 + C_ {n-1} + C_{n-2}$ pour $n \geqslant 2$.
<br/>
La suite $(u_n)_{n \in \mathbb{N}}$ définie par $u_n = C_n +1$ vérifie donc la relation de récurrence $u_{n+2} = u_{n+1} + u_n$ ; on peut montrer qu'il existe une constante $\alpha$ telle que $u_n \sim \alpha \left( \frac{1+\sqrt{5}}{2}\right)^n$, donc que $(u_n)_{n \in \mathbb{N}}$ tend rapidement vers $+\infty$.
</span>



<font size="5">🧑🏾‍💻</font>  Écrire une fonction `fibo2` telle que `fibo2 n` calcule (efficacement !) le couple $(u_n,u_{n+1})$.

<!-- grade_id: cell-9c6acc236cd06f7c, points: 0, grade: false, solution: true -->
```OCaml

let rec fibo2 n =
  match n with
  | 0 -> 0, 1
  | _ -> let a, b = fibo2 (n-1) in b, a+b
;;

```

Output:
```
val fibo2 : int -> int * int = <fun>
```

<!-- grade_id: cell-0c8bba5837a1c67f, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (fibo2 2 = (1, 2));;
let () = assert (fibo2 6 = (8, 13));;
let () = assert (fibo2 100 = (3736710778780434371, 1298777728820984005));;
(* BEGIN HIDDEN TESTS *)
let () = assert (fibo2 0 = (0, 1));;
let () = assert (fibo2 1 = (1, 1));;
let () = assert (fibo2 15 = (610, 987));;
let () = assert (fibo2 42 = (267914296, 433494437));;
(* END HIDDEN TESTS *)
```

<font size="5">👨‍💻</font> Combien d'additions faut-il effectuer pour calculer $u_n$ avec cette méthode ?

<div class="alert alert-block alert-info">
    On pourra noter $A(n)$ le nombre d'additions réalisées lors de l'appel <code>fibo2 n</code> et déterminer une relation de récurrence vérifiée par $A(n)$.
</div>

<!-- grade_id: cell-25c89611c0916d7a, points: 0, grade: true, solution: true -->



<span style='color:Teal '>
Le nombre d'addition $A(n)$ réalisées lors de l'appel `fibo2 n` vérifie $A(0) = 0$ et pour tout $n \in \mathbb{N}^*$, $A(n) = 1 + A(n-1)$.

Il s'agit donc d'une suite arithmétique de raison $1$ et de premier terme $0$, donc pour tout $n \in \mathbb{N}$, $A(n) = n$.

La complexité de la fonction `fibo2` est par conséquent linéaire.
</span>



On peut montrer que : $\forall p \geqslant 0,  \forall q \geqslant 1, u_{p+q}=u_{p+1}u_q+u_pu_{q-1}$.

<font size="5">👨🏽‍💻</font> Exprimer $u_{2n}, u_{2n+1}$ et $u_{2n+2}$ en fonction de $u_n$ et $u_{n+1}$.

<!-- grade_id: cell-6ce9a41f420e75bc, points: 0, grade: true, solution: true -->


<span style='color:Teal '>
$u_ {2n} = u_{n+1}u_n + u_nu_{n-1} = u_{n+1}u_n +u_n(u_{n+1}-u_n) = 2u_{n+1}u_n - u_n^2$.

$u_{2n+1} = u_{n + (n+1)} = u_{n+1} ^2 + u_n^2$

$u_{2n+2} = u_ {n+2}u_{n+1}+u_{n+1}u_n = (u_{n+1}+u_n)u_{n+1}+u_{n+1}u_n = u_{n+1}^2+2u_{n+1}u_n$

</span>


<font size="5">🧑‍💻</font>  En déduire une fonction `fibo3` telle que `fibo3 n` calcule le couple $(u_n,u_{n+1})$ (on distinguera deux cas suivant la parité de $n$).

<!-- grade_id: cell-bde71b7cd62722d8, points: 0, grade: false, solution: true -->
```OCaml

let rec fibo3 n =
   match n with
   | 0 -> 0, 1
   | _ ->
     let a, b = fibo3 (n/2) in
     if n mod 2 = 0
     then 2*b*a - a*a, a*a + b*b
     else a*a + b*b, b*b + 2*a*b 

```

Output:
```
val fibo3 : int -> int * int = <fun>
```

<!-- grade_id: cell-9140fd3b0b850c75, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (fibo3 2 = (1, 2));;
let () = assert (fibo3 6 = (8, 13));;
let () = assert (fibo3 100 = (3736710778780434371, 1298777728820984005));;
(* BEGIN HIDDEN TESTS *)
let () = assert (fibo3 0 = (0, 1));;
let () = assert (fibo3 1 = (1, 1));;
let () = assert (fibo3 15 = (610, 987));;
let () = assert (fibo3 42 = (267914296, 433494437));;
(* END HIDDEN TESTS *)
```

<font size="5">👩🏾‍💻</font> Évaluer le nombre d'appels à la fonction `fibo3` effectués par l'appel `fibo3 n`.

<div class="alert alert-block alert-info">
On pourra noter $u_p$ le nombre d'appels réalisés pour un entier non nul $n$ dont l'écriture binaire comporte $p$ bits.
</div>

<!-- grade_id: cell-480c88f6d36ac957, points: 0, grade: true, solution: true -->


<span style='color:Teal '>
Si $p=1$, alors $n=1$ et $u_p = 2$ (un appel avec $n=1$, qui réalise un appel avec $n=0$).

Si $p \geqslant 2$, alors $u_p = 1 + u_ {p-1}$.

Par conséquent, $u_p = p+1$.

Finalement, le nombre d'appels réalisés lors du calcul de `fibo3 n` est de l'ordre de $p$, i.e de l'ordre de $\log n$.
</span>
