# Pour démarrer

Vous vous trouvez actuellement dans un *notebook* Jupyter. Celui-ci contient deux types de cellule : des cellules de texte (comme celle-ci), et des cellules de code.

Les cellules peuvent être exécutées via le menu **Run**, via l'icône Play ou avec un raccourci clavier :
- `Shift+Entrée` pour exécuter la cellule et passer à la suivante ;
- `Ctrl+Entrée` pour exécuter la cellule sans passer à la suivante ;
- `Alt+Entrée` pour exécuter la cellule et créer une nouvelle cellule en dessous.

Par exemple, la cellule de code suivante définit une fonction permettant de calculer la factorielle d'un entier naturel.

```OCaml
let rec fact n =
    match n with
    | 0 -> 1
    | _ -> n * fact (n-1)
;;
```

La cellule ci-dessous appelle la fonction `fact` et ne pourra donc être exécutée sans erreur que si la cellule du dessus a été exécutée.

```OCaml
fact 6;;
```

Par défaut, les nouvelles cellules créées sont des cellules de code. Pour modifier le style d'une cellule, passez en mode commande à l'aide de la touche `Echap` si vous êtes en mode édition, puis tapez `m` pour transformer la cellule courante en cellule de texte ou `y` pour la transformer en cellule de code.

Vous pouvez par ailleurs créer une nouvelle cellule au-dessus de la cellule courante en tapant `A` (*above*) ou en dessous en tapant `B`.

Pour passer du mode commande au mode édition, taper sur `Entrée`.

Pour la mise en forme des cellules de texte, vous pouvez consulter l'aide sur Markdown présente dans le menu **Help**

Les cellules déjà présentes dans ce *notebook* sont de plusieurs types :
- Des cellules contenant l'énoncé, qui sont en lecture seule ;
- Des cellules de réponses, que vous pouvez (et devez) modifier ;
- Des cellules permettant de tester vos réponses, qui sont en lecture seule.

Vous pouvez ajouter d'autres cellules si vous le souhaitez.

# Pour rendre le TP

Les TP ne seront pas à rendre systématiquement. Lorsqu'un TP devra être rendu, cela sera indiqué dans le *notebook*.

**Ce premier TP est à rendre en traitant (au moins) les questions de programmation.** Les questions portant sur la complexité des fonctions sont facultatives.

Pour cela, il faudra déposer votre *notebook* au nom `TP01.ipynb` dans les deux semaines suivant le TP dans le dossier `Depot_fichiers` présent [à l'adresse suivante](https://cloud.maths.mlong.fr/), en vous connectant avec votre adresse mail et le mot de passe que vous aurez défini (utiliser *mot de passe oublié* pour la première connexion).

# Quelques fonctions autour des dates

#### Règles de validité d'une date :
- janvier, mars, mai, juillet, août, octobre, décembre comportent 31 jours ; avril, juin, septembre, novembre comportent 30 jours ; février comporte  29 jours si l'année est bissextile, 28 sinon ;
- une année est bissextile
  - pour les années séculaires, si elle est divisible par 400
  - pour les autres années, si elle est divisible par 4

<font size="5">👩‍💻</font>  Écrire la fonction Caml `bissextile : int -> bool` qui renvoie `true` si l'année `n` est bissextile.

<div class="alert alert-block alert-info">
    On pourra utiliser une expression <code>if e1 then e2 else e3</code>, où <code>e1</code> est une expression booléenne. Si la valeur de <code>e1</code> est <code>true</code>, l'expression <code>e2</code> est évaluée et sa valeur est retournée ; sinon, ce sera la valeur de <code>e3</code> qui sera retournée.</div>

<!-- grade_id: cell-342f629193548472, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-e41c9ca7fd9ba3d5, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (not (bissextile 2022));;
let () = assert (bissextile 2024);;
let () = assert (bissextile 2000);;
let () = assert (not (bissextile 2100));;
```

<font size="5">🧑🏿‍💻</font> La fonction suivante utilise un *filtrage*. La tester. Que renvoie-t-elle ?

```OCaml
let nb_jours m a =
  match m with
  | 2 -> if bissextile a then 29 else 28
  | 4 -> 30
  | 6 -> 30
  | 9 -> 30
  | 11 -> 30
  | _ -> 31
;;

```

<!-- grade_id: cell-f0208aaf427503bc, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">👨🏻‍💻</font>  Écrire une fonction `date_valide : int -> int -> int -> bool` prenant en argument les entiers `j, m , a` et renvoyant le booléen `true` si la date représentée par le triplet `j, m , a` (jour, mois, année) est une date valide, et `false` sinon.

<!-- grade_id: cell-fc2d30efe1442906, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-f1dc1281d66471ef, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (date_valide 31 07 2022);;
let () = assert (not (date_valide 29 02 2022));;
let () = assert (not (date_valide 30 14 2022));;
let () = assert (not (date_valide (-2) 14 2022));;
let () = assert (not (date_valide 18 0 2022));;
```

			
On convient de numéroter les jours de la semaine de $0$ à $6$ (0 pour dimanche, 1 pour lundi, etc.).

<font size="5">🧑🏾‍💻</font> Écrire une fonction donnant, pour tout numéro entre $0$ et $6$, le nom du jour correspondant sous forme d'une chaîne de caractères.

<div class="alert alert-block alert-info">
    On pourra utiliser un filtrage.
</div>

<!-- grade_id: cell-b73a552ab7f16ffc, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-f55cb2954afbdf55, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (nom_jour 1 = "lundi");;
let () = assert (nom_jour 2 = "mardi");;
let () = assert (nom_jour 3 = "mercredi");;
let () = assert (nom_jour 4 = "jeudi");;
let () = assert (nom_jour 5 = "vendredi");;
let () = assert (nom_jour 6 = "samedi");;
let () = assert (nom_jour 0 = "dimanche");;
```

En notant $\lfloor x \rfloor$ la partie entière de $x$, le numéro du jour de la semaine correspondant à une date donnée $(j,m,a)$ s'obtient en prenant le reste de la division par $7$ de l'entier $n$ défini par :
		$$n = j + \lfloor 2,6 p -0,2\rfloor + d + \left\lfloor \frac{d}{4} \right\rfloor +\left\lfloor \frac{c}{4} \right\rfloor +5c \qquad\text{(Formule de Zeller)}$$

où 

$$\displaystyle p = \left\lbrace \begin{array}{ll} m+10 & \text{ si } m \leqslant 2 \\
			m-2 & \text{ sinon}
		\end{array}\right. 
		\qquad b = \left\lbrace \begin{array}{ll} a-1 & \text{ si } m \leqslant 2 \\
			a & \text{ sinon}
		\end{array}\right. 
		\qquad c = \left\lfloor \frac{b}{100} \right\rfloor 
		\qquad d=b-100c $$

<font size="5">👩🏼‍💻</font>  Écrire une fonction qui calcule ce numéro à partir de $j$, $m$ et $a$.

<div class="alert alert-block alert-info">
On rappelle que OCaml est un langage fortement typé et qu'il n'est pas possible de réaliser directement une opération entre un flottant et un entier.
    
On pourra donc utiliser les fonctions <code>float_of_int</code> et <code>int_of_float</code> (qui donne la partie entière du flottant passé en paramètre).
</div>

<div class="alert alert-block alert-info">
    
Pour définir un objet dont la valeur dépend d'une condition, on pourra écrire : 
    <code>let var = if condition then e1 else e2 in</code>
</div>

<!-- grade_id: cell-78a91564a5742a3d, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-115d60a767e15dfa, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (zeller 1 1 2022 = 6);;
let () = assert (zeller 17 8 1936 = 1);;
let () = assert (zeller 23 6 1912 = 0);;
```

<font size="5">🧑🏻‍💻</font> En déduire une fonction `jour_semaine` donnant le jour de la semaine correspondant à une date donnée.

<!-- grade_id: cell-c0229bc3ee524390, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-fe5383bf1a22cf26, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (jour_semaine 10 12 1815 = "dimanche");;
let () = assert (jour_semaine 26 12 1791 = "lundi");;
```

# Calculs de puissances

<font size="5">👩🏽‍💻</font> Écrire une fonction `puiss : int -> int -> int` telle que si $n$ est un entier naturel,
`puiss x n` calcule $x^{n}$ (sans utiliser `**`).

<!-- grade_id: cell-1764cba30500d0a1, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-85acdb38f2ce94da, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (puiss 2 10 = 1024);;
let () = assert (puiss 2 11 = 2048);;
```

<font size="5">🧑‍💻</font> Donner la complexité du calcul de `puiss x n`, en nombre de multiplications effectuées, en fonction de $n$.

<!-- grade_id: cell-50bd8a271a86a887, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">👨🏽‍💻</font> Si votre implantation de `puiss` demande plus de $10$ multiplications pour calculer $2^{42}$, récrire votre fonction.

Indication: lorsque $n$ est pair, $x^{n} = \left(x^{n/2}\right)^2$ (ou, si l'on préfère $\left(x^2\right)^{n/2}$).

<!-- grade_id: cell-a3f6a061e197eb1e, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-a21724f687584e22, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (puiss 2 10 = 1024);;
let () = assert (puiss 2 11 = 2048);;
```

<font size="5">🧑🏼‍💻</font> Soit $n$ un entier naturel non nul. Quelle est le nombre de multiplications effectuées dans le pire des cas lors du calcul de `puiss x n` lorsque l'entier naturel $n$ s'écrit sur $p$ bits ?

<div class="alert alert-block alert-info">
On pourra noter $u_p$ le nombre d'additions réalisées dans le pire des cas lors de l'appel <code>puiss n</code> et déterminer une relation de récurrence vérifiée par $u_p$.
</div>

En déduire la complexité dans le pire des cas de `puiss n` en fonction de $n$.

<div class="alert alert-block alert-info">    
$p$ est l'unique entier tel que $2^{p-1} \leqslant n < 2^{p}$, donc $p = \lfloor \log_2 n \rfloor + 1$. 
</div>

<!-- grade_id: cell-0a2aeed577fb7df2, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

# Suite de Fibonacci

On considère la suite de Fibonacci $(u_n)_\mathbb{N}$ définie par  $u_0=0$, $u_1=1$ et $\forall n \in \mathbb{N}, u_{n+2}=u_{n+1}+u_n$

<font size="5">👩🏻‍💻</font>  A partir de la définition, écrire une fonction `fibo` telle que `fibo n` calcule le terme $u_n$.

<!-- grade_id: cell-869b38267042bec8, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-a0f6be5b4d426b50, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (fibo 2 = 1);;
let () = assert (fibo 10 = 55);;
```

<font size="5">👨🏿‍💻</font> Calculer quelques uns des premiers termes, puis calculer $u_{40}$. Comment expliquer que le temps de calcul soit si long ?

<!-- grade_id: cell-15ae2b5abaa63894, points: 0, grade: true, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-7836f957bc71d8c5, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">🧑🏾‍💻</font>  Écrire une fonction `fibo2` telle que `fibo2 n` calcule (efficacement !) le couple $(u_n,u_{n+1})$.

<!-- grade_id: cell-9c6acc236cd06f7c, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-0c8bba5837a1c67f, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (fibo2 2 = (1, 2));;
let () = assert (fibo2 6 = (8, 13));;
let () = assert (fibo2 100 = (3736710778780434371, 1298777728820984005));;
```

<font size="5">👨‍💻</font> Combien d'additions faut-il effectuer pour calculer $u_n$ avec cette méthode ?

<div class="alert alert-block alert-info">
    On pourra noter $A(n)$ le nombre d'additions réalisées lors de l'appel <code>fibo2 n</code> et déterminer une relation de récurrence vérifiée par $A(n)$.
</div>

<!-- grade_id: cell-25c89611c0916d7a, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

On peut montrer que : $\forall p \geqslant 0,  \forall q \geqslant 1, u_{p+q}=u_{p+1}u_q+u_pu_{q-1}$.

<font size="5">👨🏽‍💻</font> Exprimer $u_{2n}, u_{2n+1}$ et $u_{2n+2}$ en fonction de $u_n$ et $u_{n+1}$.

<!-- grade_id: cell-6ce9a41f420e75bc, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*

<font size="5">🧑‍💻</font>  En déduire une fonction `fibo3` telle que `fibo3 n` calcule le couple $(u_n,u_{n+1})$ (on distinguera deux cas suivant la parité de $n$).

<!-- grade_id: cell-bde71b7cd62722d8, points: 0, grade: false, solution: true -->
```OCaml
(*À remplacer par votre code*)
failwith "Code manquant"
```

<!-- grade_id: cell-9140fd3b0b850c75, points: 1, grade: true, solution: false -->
```OCaml
(* Exécutez cette cellule pour tester votre fonction *)
let () = assert (fibo3 2 = (1, 2));;
let () = assert (fibo3 6 = (8, 13));;
let () = assert (fibo3 100 = (3736710778780434371, 1298777728820984005));;
```

<font size="5">👩🏾‍💻</font> Évaluer le nombre d'appels à la fonction `fibo3` effectués par l'appel `fibo3 n`.

<div class="alert alert-block alert-info">
On pourra noter $u_p$ le nombre d'appels réalisés pour un entier non nul $n$ dont l'écriture binaire comporte $p$ bits.
</div>

<!-- grade_id: cell-480c88f6d36ac957, points: 0, grade: true, solution: true -->
✍️ *Votre réponse*