Si vous devez rendre votre travail, vérifiez que celui-ci peut s'exécuter sans erreur, en redémarrant le noyau et en exécutant toutes les cellules (dans le menu : Kernel$\rightarrow$Restart & Run all)

Par ailleurs, renseignez vos noms dans la cellule suivante :

```OCaml
let noms_auteurs = ["NOM_Prenom1"; "NOM_Prenom2"; "NOM_Prenom3"]
```

---

# Avant de commencer

```OCaml
let u0_test = 29818;;
```

Créer une variable `u_0` correspondant à la valeur qui vous a été attribuée.

<!-- grade_id: cell-ae4fd77f9eda066a, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

-------------------
Une entreprise de $n$ employés ($n \geqslant 2$) veut récompenser $k$ ses meilleurs vendeurs. Chaque employé, désigné par son numéro de badge $i$ (où $ 0 \leqslant i < n$), a réalisé un chiffre de ventes $a_i$ au cours de  l'année 2020. On suppose que pour tout entier $i$, $a_i$ est un entier.

L'objectif est de déterminer le $k$-ième chiffre de ventes le plus élevé, ce qui correspondra au seuil à partir duquel les employés auront le droit à une prime.


**Dans tout le sujet, on pourra utiliser ` Array.make` et ` Array.length`, mais pas d'autres fonctions du module ` Array`**

# Préliminaires

On considère une suite d'entiers $(u_k)_{k \in \mathbb{N}}$ vérifiant la relation de récurrence :
$$\forall k \in \mathbb{N}, u_{k+1} = (15091\times u_k)\text{ mod } 64007$$

#### Question 1
Écrire une fonction `suivant : int -> int` qui, lorsqu'elle prend en argument l'entier correspondant à $u_k$, renvoie $u_{k+1}$.

<!-- grade_id: cell-d95afe59143fcb79, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

#### Question 2
Écrire une fonction `tab_u : int -> int -> int array` qui prend en argument la valeur de $u_0$ et un entier $n$ et qui renvoie le tableau $[|u_0;u_1;\dots;u_{n-1}|]$ des $n$ premières valeurs de la suite $u$.

<!-- grade_id: cell-f9c992ee525430b0, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
tab_u u0_test 10;;
```

<!-- grade_id: cell-3e1361ce478079ae, points: 0, grade: true, solution: false -->
```OCaml
(*Test avec u0_test*)
let () = assert (tab_u u0_test 10 = [|29818; 14228; 35270; 41365; 42951; 38659; 43171; 30315; 25636; 14568|]);;

```

#### Question 3
Calculer les valeurs de $u_{20}$, $u_{2000}$, $u_{20000}$.

<!-- grade_id: cell-02c3613aa1d3ad63, points: 0, grade: true, solution: true -->
```OCaml
let u_20 = () ;;
let u_200 = () ;;
let u_20000 = ();;
(* ENTRER VOTRE CODE ICI *)
```

Dans toute la suite, on considère que le tableau des $n$ chiffres de ventes ($n$ pouvant dépendre de la question) est donné par ` tab_u u0 n` où ` u0` est votre valeur pour $u_0$.

# Partie I

#### Question 4
Dans cette question, on suppose que $k = 1$. 
-  Écrire une fonction ` meilleure_vente : 'a array -> 'a` qui prend en argument le tableau des chiffres de ventes et qui renvoie le meilleur chiffre de ventes.

<!-- grade_id: cell-209dddc6e211e91f, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

- Calculer la valeur de la meilleure vente pour $n = 1000$.

<!-- grade_id: cell-4af7b2e6b03f1f10, points: 0, grade: true, solution: false -->
```OCaml
(*Test avec u0_test*)
let () = assert (meilleure_vente (tab_u u0_test 1000) = 63984);;
```

<!-- grade_id: cell-bffa523d0276d073, points: 0, grade: true, solution: true -->
```OCaml
let meilleure_vente_1000 = ();;
(* ENTRER VOTRE CODE ICI *)
```

#### Question 5
Dans cette question, on suppose que la distribution des chiffres de ventes est relativement uniforme. Une approximation du seuil de ventes à atteindre pour obtenir la prime est alors donnée par la moyenne du chiffre de ventes minimum et du chiffre de ventes maximum, respectivement pondérés par $(k-1)$ et $n-k$.

- Écrire une fonction ` approx_seuil: int array -> int -> int`  prenant en argument le tableau des chiffres de ventes et l'entier $k$ et renvoyant cette moyenne pondérée arrondie à l'entier inférieur. 

<!-- grade_id: cell-7926e6e3546b6e36, points: 0, grade: false, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

- Calculer le seuil obtenu avec la fonction précédente pour $k = 200$ et $n=5000$.

<!-- grade_id: cell-0c48a3f8ca40c47e, points: 0, grade: true, solution: false -->
```OCaml
(*Test avec u0_test*)
let () = assert (approx_seuil (tab_u u0_test 5000) 200 = 61455)
```

<!-- grade_id: cell-78cd40ef46651c51, points: 0, grade: true, solution: true -->
```OCaml
let seuil_approx_200_5000 = ();;
(* ENTRER VOTRE CODE ICI *)
```

# Partie II
On traite maintenant le cas général, en ne faisant aucune hypothèse sur la répartition des ventes.

Pour calculer le $k$-ième meilleur chiffre de ventes, on va trier partiellement le tableau dans l'ordre décroissant, de sorte que les $k$ plus grands éléments soient à leur place dans le tableau.

#### Question 6
Écrire une fonction `echange : 'a array -> int -> int -> unit` qui prend en argument un tableau et deux entiers $i$ et $j$ et échangeant en place les éléments d'indices $i$ et $j$ du tableau. Aucune vérification sur $i$ et $j$ n'est demandée.

<!-- grade_id: cell-a90ef73f0d941cf1, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

#### Question 7 
Écrire une fonction ` seuil : 'a array -> int -> 'a` prenant en argument le tableau des chiffres de ventes et l'entier $k$ et renvoyant le $k$-ième plus grand élément du tableau.

<!-- grade_id: cell-a4ef26d13b55a7fe, points: 0, grade: false, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

#### Question 8
Calculer le seuil obtenu avec la fonction précédente pour $k = 200$ et $n=5000$.

<!-- grade_id: cell-322c0828bea48beb, points: 0, grade: true, solution: false -->
```OCaml
(*Test avec u0_test*)
let () = assert (seuil (tab_u u0_test 5000) 200 = 61792)
```

<!-- grade_id: cell-504d7f92986b6521, points: 0, grade: true, solution: true -->
```OCaml
let seuil_II_200_5000 = ();;
(* ENTRER VOTRE CODE ICI *)
```

# Partie III
Le nombre d'employés $n$ étant particulièrement grand, on cherche à optimiser le calcul du $k$-ième meilleur chiffre de ventes.

Pour cela, on procède comme suit : on choisit un employé $i$, et on place son chiffre de ventes $v = a_i$ à l'indice $l$ qu'il aurait dans le tableau si celui-ci était trié dans l'ordre décroissant, en réordonnant le tableau de sorte que $a_j > v$ si $j < l$ et $a_j \leqslant v$ et $j>l$. Si l'indice $l$ correspond au $k$-ième élément du tableau, on a terminé, sinon, on recommence avec la partie à gauche ou à droite de $l$ dans le tableau.

#### Question 9
Soient $g$ et $d$ deux entiers tels que $0 \leqslant g < d < n$ et soit $v = a_g$. Écrire une fonction ` partition : 'a array -> int -> int -> int` qui prend en argument le tableau des chiffres de ventes et les entiers $g$ et $d$ (dans cet ordre), qui réordonne le tableau en place entre les indices $g$ et $d$ compris de sorte que $a_j > v$ pour $g \leqslant j < l$ et $a_j \leqslant v$ pour $l < j \leqslant d$, et qui renvoie l'indice $l$ de la valeur $v$ après modification du tableau.

*Indication :* On procédera par des échanges ; on pourra laisser $v$ à l'indice $g$ jusqu'à avoir fini de parcourir la portion de tableau puis procéder à un dernier échange avec l'indice $l$.

<!-- grade_id: cell-d1edc2c045057cec, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

#### Question 10
Calculer l'indice obtenu avec la fonction précédente pour $n = 5000$, $g = 1000$ et $d = 2000$.

<!-- grade_id: cell-70658e06a6608ea4, points: 0, grade: true, solution: false -->
```OCaml
(*Test avec u0_test*)
let () = assert (partition (tab_u u0_test 5000) 1000 2000 = 1926)
```

<!-- grade_id: cell-49d0ec9da891b2f3, points: 0, grade: true, solution: true -->
```OCaml
let indice_5000_1000_2000 = ();;
(* ENTRER VOTRE CODE ICI *)
```

#### Question 11
Écrire une fonction récursive ` seuil_rec : 'a array -> int -> int -> int -> 'a` qui prend en argument le tableau des chiffres de ventes, trois entiers $j$, $g$ et $d$ (dans cet ordre) tels que $0 \leqslant g \leqslant j \leqslant d$ et qui renvoie l'élément qui serait à l'indice $j$ dans le tableau des ventes trié dans l'ordre décroissant. On pourra utiliser la fonction ` partition`.

<!-- grade_id: cell-4cb42b7df2de9ba0, points: 0, grade: true, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

#### Question 12
En déduire une fonction `seuil_bis : 'a array -> int -> 'a` prenant en argument le tableau des chiffres de ventes et l'entier $k$ et renvoyant le $k$-ième plus grand élément du tableau.

<!-- grade_id: cell-26c5f5e6d7fd51ad, points: 0, grade: false, solution: true -->
```OCaml
(* ENTRER VOTRE CODE ICI *)
```

#### Question 13
Calculer le seuil obtenu avec la fonction précédente :
+ pour $k = 1000$ et $n=20000$ ;
+ pour $k = 10000$ et $n=500000$.

<!-- grade_id: cell-f0c459e53325c546, points: 0, grade: true, solution: false -->
```OCaml
(*Test avec u0_test*)
let () = assert (seuil_bis (tab_u u0_test 20000) 1000 = 60826)
let () = assert (seuil_bis (tab_u u0_test 500000) 10000 = 62725)
```

<!-- grade_id: cell-7ba0933deb63d1d4, points: 0, grade: true, solution: true -->
```OCaml
let seuil_III_1000_20000 = ();;
let seuil_III_10000_500000 = ();;
(* ENTRER VOTRE CODE ICI *)
```